# Lab 6—“Kind of Blue”

<table>
<tr><td>
<img src="http://www.benmichel.com/wp-content/uploads/2015/01/BLINDWILLIEJOHNSONimg.jpg" />
</td><td>
For this lab project, we are going to work with a list of blues musicians in various traditions.  You should have the following files already in your `lab6` directory:
    
    lab6.txt
    process.py
    blues.txt
    zydeco.txt

Open the file `blues.txt` in a text editor and examine its format.
</td>
</tr>
</table>
    # Delta Blues
    Cecil Augusta
    Mose Allison
    Tommy Bankhead
    ...
    Elder Roma Wilson
    
    # Chicago Blues
    Alberta Adams
    Luther Allison
    ...

We are going to `read` this file and run it through the standard pipeline of data analysis:  data cleaning and preparation, data processing, and output.


If you have a line of text as a string:

    example = 'Jimmy Clanton'
    
and you'd like to turn it into a tuple `record = (surname, first_name)`, how would you do it?

First, you'd probably want to split it into pieces so there is a separate `first_name` and `surname` to assign:

    names = example.split(' ')

Now you have a `list` called `names` which looks like this:

    ['Jimmy', 'Clanton']

which you can use to assign the separate variables in your tuple either in two steps:

    surname = names[1]
    first_name = names[0]
    record = (surname, first_name)

or in one step:

    record = (names[1], names[0])

This gets a little trickier with multiple first names:

    example = 'John Henry Barbee'

In this case, I suggest the following method.  First, split the name into its pieces like this:

```python
names = example.split(' ')
```

<!-- grade_id: process, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: process-test-returntype, points: 1, grade: true, solution: false -->
```python
zydeco = process('zydeco.txt')
assert type(zydeco) == list
```

<!-- grade_id: process-test-returnvalues, points: 1, grade: true, solution: false -->
```python
zydeco = process('zydeco.txt')
assert zydeco[0] == ('Chavis', 'Boozoo', 'Zydeco')
assert zydeco[-3] == ('Dopsy', "Rockin'", 'Zydeco')
```